# Welcome to Lab_Plots!

You may not know this, but the salaries of the employees at the University of Illinois are publicly available! We have curated this data in a nice, structured dataset called `graybook.csv` for you to explore.  The goal of this lab is to work with real UIUC salary data to explore its properties, answer important questions, and to think about the implications of collecting and analyzing this data.  Throughout the lab, it is important to think about being a critical consumer of data who can not only use statistics and programming to analyze data but can also think about the "why" part of data science both in the classroom and in the world. Let’s get started!

In this lab, you will:

- Work with real UIUC salary data to explore some of the statistics that we talked about in class: mean, median, standard deviation, etc.
- Practice creating plots to visualize quantitative data: boxplots and histograms.
- See how data science can be used in the real world to think about important issues through individual reflection and discussions with your group.

## PRE-LAB SURVEY

We are doing a pre and post lab survey this week! Please complete this survey BEFORE you start the lab to get +2 extra credit points! Who doesn't want extra credit?!  This will help us make the labs better and should take about 5 minutes :) Click the link below to access the survey.

[https://forms.gle/ouMaFuucdJHXKVza7](https://forms.gle/ouMaFuucdJHXKVza7)

## 000. Your Group

Edit the next Python cell to add information about who you're working within your lab section:

```python
partner1_name = ""
partner1_netid = ""
partner1_major = ""

partner2_name = ""
partner2_netid = ""
partner2_major = ""

partner3_name = ""
partner3_netid = ""
partner3_major = ""
```

## 00. Install MatPlotLib

In this lab, we will be creating two types of simple visualizations: Histograms and Box Plots.  To do this, we need to make sure `matplotlib` is installed.

Pull up your terminal and run the following code:

MAC USERS: 

python3 -mpip install matplotlib


WINDOWS USERS: 

py -m pip install matplotlib



## 0. Import the Data

In the same directory as this notebook, find the dataset called `graybook.csv` that is provided.  We will use the other dataset later on ;) You know what to do to import the libraries and read in the data! :)

Look at the variable names, how many rows and columns there, and anything else you may notice!

<!-- grade_id: cell-27971beecac131cc, points: 0, grade: false, solution: true -->
```python
#import the graybook.csv dataset below 
```

**❓ Individual Reflection and Group Discussion ❓**

Thinking about the variables in the dataset, what are two questions that you’d like to use data science to answer (ex. What is the mean salary of my home department?).  Type them below and then share at least one of your questions with your group and why you want to find the answer to it.

Double Click this cell to edit it!

Question 1:

Question 2:

## 1. Exploratory Data Analysis (EDA)
As we discussed in lecture, the first step of any data analysis is to get familiar with your dataset.  Think about what this data can tell you and what variables are included.  Data scientists always start with this step.


Let’s do some overall exploratory data analysis to get some baseline data to compare to.  Before you do the calculations, guess what you think the average salary is at UIUC. Enter your answer below.  

Guess for Overall Average Salary: double click to replace this text with your guess

### Puzzle 1.1: Descriptive Statistics

Next, using all of the data, find the following information:

1. the number of employees at UIUC
2. the mean salary
3. the median salary
4. the standard deviation

<!-- grade_id: cell-419796a0b415e588, points: 0, grade: false, solution: true -->
```python
university_num_emp = ...
university_num_emp


```

<!-- grade_id: cell-4dbbeef1022a20e4, points: 0, grade: false, solution: true -->
```python
university_mean_sal = ...
university_mean_sal


```

<!-- grade_id: cell-727195c3fef2cf78, points: 0, grade: false, solution: true -->
```python
university_median_sal = ...
university_median_sal


```

<!-- grade_id: cell-35e651d14a628e50, points: 0, grade: false, solution: true -->
```python
university_std_sal = ...
university_std_sal


```

<!-- grade_id: cell-04e624817d681f8e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
x = university_num_emp + university_mean_sal + university_median_sal + university_std_sal
y = ((university_mean_sal - university_median_sal) * university_std_sal) / university_num_emp

assert( abs(x - 246964.21165107167) < 1 ), "It looks like your calculations are off?"
assert( abs(y - 221601.415478511) < 1 ), "It looks like your calculations are off?"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

**❓ Individual Reflection ❓**

Write a few sentences answering the following questions: Is the mean or median larger? Why do you think this might be the case? What does the standard deviation tell you in this context?

Double Click this cell to edit and type your answer to the individual reflection question!

### Visual Displays of Data -- A Key Part of EDA!

Now, we are a bit more familiar with the dataset and some summary statistics. Looking at overall descriptive statistics helps us summarize all of the observations in a column, rather than having to scroll through all of the observations!  However, descriptive statistics alone often don’t tell the whole story.  

In lecture, we will see Anscombe’s Quartet (a set of 4 datasets with the exact same descriptive statistics that look completely different when graphed). This shows us that visual displays of data are also a necessary and important part of exploratory data analysis. 

### Puzzle 1.2: Histograms

Let’s start by looking at one of the most simple, yet powerful visual displays of data, a histogram.  Histograms are used to show the overall shape of the data and they allow us to see frequencies.  Create a frequency histogram of the salaries at UIUC.


```python
# create your histogram here
```

### Puzzle 1.3: Boxplots

Next, let’s look at another simple, yet power visualization: a boxplot! Create a boxplot of the overall salary data at UIUC. 

```python
# create your boxplot here
```

**❓ Group Discussion ❓**

Discuss with your group whether you think a histogram or a boxplot or both best visualize the salary data.  There is no right answer to this question.  Explain why histograms, boxplots, or both are important and what they can tell us about the data.  

Why is it valuable to look at visual displays of salary data in general (as opposed to just looking at descriptive statistics like the mean and SD)?

## 2. Department-Level Data

One of the interesting properties of the salary dataset is that you can see what department each professor is from and look at the same data at a department level.  In other words, you can calculate descriptive statistics and look at visual displays of data for each department.  Let’s explore whether or not there are salary discrepancies among a few departments.

### Puzzle 2.1
There was a certain function you learned that will help you **group** all of the people in each department so you can find departmental aggregates.  In other words, you want to group by `'Primary Department'`.

Next, instead of just grouping by a single aggregate, use **Python list syntax** to aggregate by `count`, `mean`, and `median` to find the mean and median salary, as well as the count for how many people are in each department.  

*(You can view technical docs here:
https://pandas.pydata.org/pandas-docs/stable/reference/groupby.html)*

<!-- grade_id: cell-72189b308296916b, points: 0, grade: false, solution: true -->
```python
group = ...
agg_df = ...


# Here, we only care about two columns so we are selecting those columns to be in our dataset
agg_df = agg_df[["Primary Department", "Salary"]]

# Display 5 random rows to make sure it works:
agg_df.sample(5)
```

<!-- grade_id: cell-d4a69c902629304a, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
assert( len(agg_df) == 140 ), "You should have 140 different departments.  Check to make sure you grouped the data?"
assert( "Salary" in agg_df ), "You must have a salary output"
assert( "mean" in agg_df["Salary"] ), "Make sure you have a mean in your salary"
assert( "count" in agg_df["Salary"] ), "Make sure you have a count in your salary"
assert( "median" in agg_df["Salary"] ), "Make sure you have a median in your salary"
assert( len(agg_df.columns) == 4 ), "Make sure you have ONLY the salary data and the department"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 2.2 

Let’s look at 3 different departments: The Department of English, The Department of Psychology, and The Department of Electrical & Computer Eng.  Find the mean and median salary for each of these departments.  We'll do the means first!

```python
# Find the means

mean_English = ...
print(mean_English)

mean_Psych = ...
print(mean_Psych)

mean_ECE = ...
print(mean_ECE)
```

```python
## == TEST CASES for Puzzle 2.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
x = mean_English + mean_Psych + mean_ECE

assert( x > 336499 ), "It looks like your calculations are off?"
assert( x < 336500 ), "It looks like your calculations are off?"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 2.3 

Find the median salaries of the following departments: The Department of English, The Department of Psychology, and The Department of Electrical & Computer Eng.  

```python
# Find the medians

median_English = ...
print(median_English)

median_Psych = ...
print(median_Psych)

median_ECE = ...
print(median_ECE)
```

```python
## == TEST CASES for Puzzle 2.3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
y = median_English + median_Psych + median_ECE

assert( y > 308547 ), "It looks like your calculations are off?"
assert( y < 308548 ), "It looks like your calculations are off?"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

**❓ Individual Reflection❓** 

Are the means and medians drastically different from each other in each department and between each of the 3 departments? Think about why or why not. Are there any confounding variables that could be driving these differences?  Write a few sentences below.  

Double Click this cell to edit and type your answer the individual reflection question!

### Puzzle 2.4

Create a histogram of the salaries in your home department.  In other words, if your major is Statistics, you should have a histogram of all the salaries in the Statistics Department. Also, find the maximum and minimum salary in your home department. If you're undecided, pick a department you are interested in.

```python
# histogram of your home department
```

```python
#find the max and min salary for your home deparment

max_salary = ...
print(max_salary)

min_salary = ...
print(min_salary)
```

**❓ Group Discussion ❓** 

Share the results with the people at your table. Do you notice any similarities or differences between departments?

## 3. Genders and Salaries

Data can often reveal systemic problems or discrimination. For example, in many companies, men and women are promoted at different rates.  Let’s look at a subset of the salary dataset to investigate whether or not there is a difference in salaries between faculty who identify as men and women in two departments. We have this data for the STAT and CS departments (Karle and Wade’s home departments) and have compiled it in a dataset called `statcssalaries.csv`. This data is also located in the same directory as this lab!

### Puzzle 3.1

Load the second dataset `statcssalaries.csv` and create two boxplots of the male salaries and female salaries in this dataset.  

<!-- grade_id: cell-dca5763613817466, points: 0, grade: false, solution: true -->
```python
# load the dataset



# create the two boxplots
```

### Puzzle 3.2

Next, calculate the overall mean, median, and SD of the salaries for this data.

<!-- grade_id: cell-8d360389fd08fc8f, points: 0, grade: false, solution: true -->
```python
overall_mean = ...
print(overall_mean)

overall_median = ...
print(overall_median)

overall_SD = ...
print(overall_SD)
```

```python
## == TEST CASES for Puzzle 3.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs with a message (with the emoji showing), you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
x = overall_median + overall_median + overall_SD
y = ((overall_mean - overall_median) * overall_SD) / overall_SD

assert( abs(x - 271945.50619890203) > 1 ), "It looks like your calculations are off?"
assert( (y - 10038.359022556397) < 1 ), "It looks like your calculations are off?"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

**❓ Group Discussion ❓**

Discuss the results with your group.  Some questions to think about: Do the boxplots look similar or different? Are there any outliers? Do you think these results present a problem?  What are the implications of this analysis for faculty, students, and the university?

**❓ Individual Reflection❓**

As a student, do you think these are fair salaries or are they too low or too high?  Why do you think this?  How does this dataset impact you? Write down your thoughts below.

Double Click this cell to edit and type your answer to the individual reflection question!

## 4. Exploring Your Own Interests

At this point of the lab, we have investigated a lot of questions, however, these have been questions that we have told you to answer. As a data scientist, it is important to be able to use the data science skills that you learn in the classroom to answer real questions that you have.  

**❓ Individual Reflection ❓**

Think about two questions that you have that have not been answered. Record them below.  Then, answer at least one of these questions using Python and either dataset.  

Double Click this cell to edit!

Unanswered Question 1:

Unanswered Question 2: 

```python
# answer a question that you have using Python in this cell- it can be something simple!
```

**❓ Group Discussion ❓**

Share your question and results with your group. 

## 5. Counter Data

As a data scientist, it is important to think about the implications of collecting and not collecting certain data.  Data that is not collected in an analysis is called *counter data*.  

Many times, data is collected to make a profit.  Data has even been referred to as “the new oil” because it is oftentimes seen as an uncapped natural resource and if you can figure out how to capture and refine it, it can lead to a massive profit.

**❓ Group Discussion ❓**
Chat in your groups about the following questions: What is some counter data that you’d like to collect in regard to salaries? Why do you think it’s important to collect this data?  Share your thoughts with your group.  Some examples of questions that would need counter data are listed below:  

- Do smaller colleges pay a lot less than big research universities? 
- How does UIUC compare to other similar schools in terms of what they pay their faculty? 
- Are there salary disparities if we stratify based on ethnicity? 
- How has this data changed over time?


**❓ Individual Reflection ❓**

Now think about how you would actually go about collecting this data to answer a question that you have.  Respond below with the question you’d want to answer, the counter data you’d want to collect, what resources you’d need to collect it, and if this is realistic. 


Double Click this cell to edit and type your answer to the individual reflection question!

## 6. Beyond Stat 107

As a Data Scientist, it’s important to think about who data analysis can benefit and who it can harm.  It’s also important to think about why the data is collected and why counter data is not collected.  

**❓ Group Discussion ❓**

Who benefits from collecting this salary data?  Who does this data harm?
Why do you think this salary data is public? Can you think of any reasons that this could be problematic? 


**❓ Last Individual Reflection ❓**

Let’s think beyond the salary data!  Thinking about the data science skills that you have mastered: experimental design, descriptive statistics, and basic visualization, what is one real world question that you’d like to answer using data?  It can be about anything, a hobby or a specific passion or issue.  You won’t be doing this, but instead thinking about how you’d do this. Write down your thoughts below.

Double Click this cell to edit!

## POST-LAB SURVEY

Remember, we are doing a pre and post lab survey this week! Please complete this survey AFTER you finish the lab to get +2 additional extra credit points! This will also help us make the labs better and should take about 10 minutes :) Click the link below to access the survey.

[https://forms.gle/5E33LkR24KgNoo6A8](https://forms.gle/5E33LkR24KgNoo6A8)

## Submit Your Work

You're almost done -- congratulations! 🎉

You need to do two more things:

1.  Save your work. To do this, go to File -> Save All

2.  After you have saved, exit this notebook and follow the webpage instructions to commit this lab to your Git repository!